# Welcome to your new dev environment!

This is a quick tester file to ensure that we can open and run Jupyter notebook files.

Fill out the hours and collaborators, then hit "Run All" to generate the appropriate fields. 


```python
# Initialize Otter
import otter
grader = otter.Notebook("environment_tester.ipynb")

# This check is currently empty and is guaranteed to pass
grader.check("tests")
```

Output:
```
tests results: All test cases passed!
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. 

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"myself"}
# List of URLS (creates a set)
websites = {"None"}
# Approximate number of hours, including lab/in-class time
hours = 2.0

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

Output:
```
hours_collaborators results: All test cases passed!
```

## Submission

Now, you'll need to **save this file** and then run the command in a terminal:

`jupyter nbconvert --to pdf environment_tester.ipynb`

Check to see that the PDF file contains some text and upload to Gradescope!

Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for list (points: 1)
```
>>> res = calc_stats_from_list([-1, 2, -3, 4, -5])
>>> assert np.isclose(res["Mean positive"], 3)
```
```
>>> res = calc_stats_from_list([-1, 2, -3, 4, -5])
>>> assert np.isclose(res["Mean negative"], -3)
```
```
>>> res = calc_stats_from_list([-1, 2, -3, 4, -5])
>>> assert np.isclose(res["Count positive"], 2)
```
```
>>> res = calc_stats_from_list([-1, 2, -3, 4, -5])
>>> assert np.isclose(res["Count negative"], 3)
```


Tests for nparray (points: 1)
```
>>> test_data = np.array([-1, 2, -3, 4, -5])
>>> res = calc_stats_from_nparray(test_data)
>>> assert np.isclose(res["Mean positive"], 3)
```
```
>>> test_data = np.array([-1, 2, -3, 4, -5])
>>> res = calc_stats_from_nparray(test_data)
>>> assert np.isclose(res["Mean negative"], -3)
>>>
```
```
>>> test_data = np.array([-1, 2, -3, 4, -5])
>>> res = calc_stats_from_nparray(test_data)
>>> assert np.isclose(res["Count positive"], 2)
```
```
>>> test_data = np.array([-1, 2, -3, 4, -5])
>>> res = calc_stats_from_nparray(test_data)
>>> assert np.isclose(res["Count negative"], 3)
```


Tests for tests (points: 2)
```

```
